<a href="https://colab.research.google.com/github/slipstreamerVI254/python-programming/blob/main/Copy_of_Assignment2_Functions.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 04.1 (simple function)

Write a function called `is_even` which takes an integer as an argument and returns `True` if the argument is even, and otherwise returns `False`. Test your function for several values.

<!-- grade_id: cell-f07e8ac1f9838913, points: 0, grade: false, solution: true -->
```python
def is_even(x):
    x = int(x)
    if x % 2 == 0:
        return True
    else:
        return False
    return

```

<!-- grade_id: cell-8a4ca795d9fa96dd, points: 0, grade: true, solution: false -->
```python
print(is_even(0))   # Output: True
print(is_even(101)) # Output: False
print(is_even(982)) # Output: True
print(is_even(-5))  # Output: False
print(is_even(-8))  # Output: True
```

Output:
```
True
False
True
False
True
```

## Exercise 04.2 (functions and default arguments)

Write a single function named `magnitude` that takes each component of a vector of length 2 or 3 and returns the magnitude.
Use default arguments to handle vectors of length 2 or 3 with the same code. Test your function for correctness against hand calculations for a selection of values.

<!-- grade_id: cell-693a09b45c6bb532, points: 0, grade: false, solution: true -->
```python
import math
def magnitude(x, y, z=0):
    return math.sqrt(x**2 + y**2 + z**2)
```

<!-- grade_id: cell-384a85ba97b226d7, points: 0, grade: true, solution: false -->
```python
## tests ##
assert math.isclose(magnitude(3, 4), 5.0)
assert math.isclose(magnitude(4, 3), 5.0)
assert math.isclose(magnitude(4, 3, 0.0), 5.0)
assert math.isclose(magnitude(4, 0.0, 3.0), 5.0)
assert math.isclose(magnitude(3, 4, 4), 6.403124237)
```

## Exercise 04.3 (functions)

Given the coordinates of the vertices of a triangle, $(x_0, y_0)$, $(x_1, y_1)$ and $(x_2, y_2)$, the area $A$ of the triangle is given by:
$$
A = \left| \frac{x_0(y_1  - y_2) + x_1(y_2 - y_0) + x_2(y_0 - y_1)}{2} \right|
$$
Write a function named `area` that computes the area of a triangle given the coordinates of the vertices.
The order of the function arguments must be (`x0, y0, x1, y1, x2, y2)`.

Test the output of your function against some known solutions.

<!-- grade_id: cell-ed02439d4509b2c1, points: 0, grade: false, solution: true -->
```python
def area(x0, y0, x1, y1, x2, y2):
    return abs((x0 * (y1 - y2) + x1 * (y2 - y0) + x2 * (y0 - y1)) / 2)
```

<!-- grade_id: cell-64f8b6708595aed3, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
x0, y0 = 0.0, 0.0
x1, y1 = 0.0, 2.0
x2, y2 = 3.0, 0.0
A = area(x0, y0, x1, y1, x2, y2)
assert math.isclose(A, 3.0)
```

## Exercise 04.4 (recursion)

The factorial of a non-negative integer $n$ is expressed recursively by:
$$
n! =
\begin{cases}
1 & n = 0 \\
(n - 1)! \,n & n > 0
\end{cases}
$$

Develop a function named `factorial` for computing the factorial using recursion.
Test your function against the `math.factorial` function, e.g.

```python
def factorial(n):
    if n == 0:
        return 1
    else:
        return n * factorial(n - 1)
```

<!-- grade_id: cell-9591c1fc1c49c167, points: 0, grade: false, solution: true -->
```python
import math
print("Reference value of factorial of 5:", math.factorial(5))
```

Output:
```
Reference value of factorial of 5: 120
```

<!-- grade_id: cell-ffe589e9fb2f1419, points: 0, grade: true, solution: false -->
```python
## tests ##
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(2) == 2
assert factorial(5) == 120

import math
assert factorial(32) == math.factorial(32)
```

## Exercise 04.5 (functions and passing functions as arguments)

Restructure your program from the bisection problem in Exercise 02 to

- Use a Python function to evaluate the mathematical function $f$ that we want to find the root of;

and then

- Encapsulate the bisection algorithm inside a Python function, which takes as arguments:
  1. the function we want to find the roots of
  1. the points $x_{0}$ and $x_{1}$ between which we want to search for a root
  1. the tolerance for exiting the bisection algorithm (exit when $|f(x)| < \text{tol}$)
  1. maximum number of iterations (the algorithm should exit once this limit is reached)

For the first step, create a Python function for evaluating $f$, e.g.:z
```python
def f(x):
    # Put body of the function f(x) here, returning the function value
```           
For the second step, encapsulate the bisection algorithm in a function:
```python
def compute_root(f, x0, x1, tol, max_it):

    # Implement bisection algorithm here, and return when tolerance is satisfied or
    # number of iterations exceeds max_it

    # Return the approximate root, value of f(x) and the number of iterations
    return x, f, num_it

# Compute approximate root of the function f
x, f_x, num_it = compute_root(f, x0=0, x1=1, tol=1.0e-6, max_it=1000)
```

You can try testing your program for a function $f(x)$ that is simpler from the function Exercise 02, e.g. $f(x) = x - 4$.
A quadratic function, the roots of which you can find analytically, would be a good test case.

### Solution

Define the function for computing $f(x)$:

```python
def my_f(x):
    """Evaluate polynomial function"""
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7
```

<!-- grade_id: cell-622edcb129dcc7f2, points: 0, grade: false, solution: true -->
```python
def compute_root(f, x0, x1, tol, max_it):
    num_iterations = 0
    f_mid = float('inf')  # Initialize f_mid to a large value
    while abs(f_mid) > tol and num_iterations < max_it:
        x_mid = (x0 + x1) / 2
        f_mid = f(x_mid)

        # The condition was incorrect in the original code
        # The if condition should check for the interval where the root lies
        if f(x0) * f_mid < 0:
            x1 = x_mid
        else:
            x0 = x_mid
        num_iterations += 1
    return x_mid, f_mid, num_iterations


```

<!-- grade_id: cell-cb5a5cdf1d790edb, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=1000)

# Test solution for function in Exercise 02
assert math.isclose(x, 1.1568354368209839)
```

#### Optional extension

Use recursion to write a `compute_root` function that *does not* require a `for` or `while` loop.

<!-- grade_id: cell-c944bb2229317d4d, points: 0, grade: false, solution: true -->
```python
def compute_root(f, x0, x1, tol, max_it, it_count=0):
    x_mid = (x0 + x1) / 2
    f_mid = f(x_mid)

    if abs(f_mid) <= tol or it_count >= max_it:
        return x_mid, f_mid, it_count  # Base case: tolerance or max iterations reached

    it_count += 1

    if f(x0) * f_mid < 0:
        return compute_root(f, x0, x_mid, tol, max_it, it_count) # Recursive call
    else:
        return compute_root(f, x_mid, x1, tol, max_it, it_count) # Recursive call

```

<!-- grade_id: cell-9ebe2f13b2899a8f, points: 0, grade: true, solution: false -->
```python
## tests ##
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=1000)
assert math.isclose(x, 1.1568354368209839)
```